# Assignment A1: Power System Data

## 🚀 Start Assignment in Binder

[![Launch Binder](https://mybinder.org/badge_logo.svg)](
https://mybinder.org/v2/gh/nilsjakob/power_systems_intro-book/main?urlpath=lab/tree/01_introduction/A1_power_system_data.ipynb
)

---

## How to work on this assignment

1. Click the **Launch Binder** button above  
2. Wait for the environment to start (30–60 seconds)  
3. Solve the tasks  
4. Run the test/check cells  
5. Repeat until everything is correct 
# Assignment A0: Power System Data (Norway)

This assignment is based on real data from the Norwegian power system.

## Learning objectives

- Interpret system-level energy data (TWh)
- Perform basic calculations
- Understand energy mix and surplus
- Reflect on system implications



## Given Data (2025)

| Variable              | Value                  |
|----------------------|------------------------|
| Production           | 162 TWh               |
| Consumption          | 139 TWh               |
| Hydropower share     | 90%                   |
| Wind power share     | 9%                    |
| Solar share          | 0.2%                  |
| Thermal share        | 1%                    |

```python
score = 0
total = 0

def check_answer(student, expected, name, hint=""):
    global score, total
    total += 1
    
    if student is None:
        print(f"{name}: no answer yet")
    elif abs(student - expected) < 1e-6:
        print(f"{name} is correct")
        score += 1
    else:
        print(f"{name} is incorrect")
        if hint:
            print(f"Hint: {hint}")

```

```python

# Data provided
production = 162
consumption = 139

hydro_share = 0.90
wind_share = 0.09
solar_share = 0.002
thermal_share = 0.01

```

## Q1: Power balance

Compute the power surplus:

Surplus = Production − Consumption

Return the result in TWh.

<!-- grade_id: cell-a2ec382ac0152dba, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
surplus = None

```

<!-- grade_id: cell-ac0598a1ecac7ae3, points: 0, grade: false, solution: true -->
```python
# SOLUTION
surplus = production - consumption
surplus


```

Output:
```
23
```

<!-- grade_id: cell-2bea0b2d6066e1ac, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
assert abs(surplus - (production - consumption)) < 1e-6
### END HIDDEN TEST


```


## Q2: Renewable share

Hydropower and wind power are considered renewable.

Compute the total renewable share (as a fraction)


<!-- grade_id: cell-ba6378cf08d08153, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
renewable_share = None

```

<!-- grade_id: cell-4f14145eeaec0d93, points: 0, grade: false, solution: true -->
```python
# SOLUTION
renewable_share = hydro_share + wind_share
renewable_share
```

Output:
```
0.99
```

<!-- grade_id: cell-671953f6688ec0f5, points: 1, grade: true, solution: false -->
```python
# TEST
expected = hydro_share + wind_share
assert abs(renewable_share - expected) < 1e-6


```


## Q3: Hydropower production

Compute the total hydropower production in TWh.


<!-- grade_id: cell-3858165536cdc21f, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
hydro = None

```

<!-- grade_id: cell-5e9ca4a73b9f31f9, points: 0, grade: false, solution: true -->
```python
# SOLUTION
hydro = production * hydro_share
hydro

```

Output:
```
145.8
```

<!-- grade_id: cell-d2e23ecc95a12f09, points: 1, grade: true, solution: false -->
```python
# TEST
expected = production * hydro_share
assert abs(hydro - expected) < 1e-6

```

## Q4: Non-renewable share

Estimate the non-renewable share of the system.

Hint: Sum all renewable sources first.


<!-- grade_id: cell-20631483efaebbe6, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
renewable = None

```

<!-- grade_id: cell-889c3b1dbf4fd685, points: 0, grade: false, solution: true -->
```python
# SOLUTION
renewable = hydro_share + wind_share + solar_share
non_renewable = 1 - renewable
non_renewable

```

Output:
```
0.008000000000000007
```

<!-- grade_id: cell-758e5b4035aa39b2, points: 1, grade: true, solution: false -->
```python
# TEST
expected = 1 - (hydro_share + wind_share + solar_share)
assert abs(non_renewable - expected) < 1e-6

```


## Q5: Wind power production

Compute wind power production in TWh.


<!-- grade_id: cell-93f32f5426992dbe, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
wind = None
```

<!-- grade_id: cell-8a31a2962176373c, points: 0, grade: false, solution: true -->
```python
# SOLUTION
wind = production * wind_share
wind

```

Output:
```
14.58
```

<!-- grade_id: cell-c295406f96fdb8ae, points: 1, grade: true, solution: false -->
```python
# TEST
expected = production * wind_share
assert abs(expected - expected) < 1e-6

```

```python

print("Checking all answers...\n")

check_answer(surplus, production - consumption, "Q1 surplus")
check_answer(renewable_share, hydro_share + wind_share, "Q2 renewable share")
check_answer(hydro, production * hydro_share, "Q3 hydro")
check_answer(non_renewable, 1-renewable, "Q4 non-renewable share")
check_answer(wind, production * wind_share, "Q5 wind power")
```

Output:
```
Checking all answers...

Q1 surplus is correct
Q2 renewable share is correct
Q3 hydro is correct
Q4 non-renewable share is correct
Q5 wind power is correct
```

```python
print("🔍 Checking all answers...\n")

check_answer(surplus, production - consumption, 
             "Q1 surplus", 
             "Subtract consumption from production")

check_answer(renewable_share, hydro_share + wind_share, 
             "Q2 renewable share", 
             "Add hydro and wind shares")

check_answer(hydro, production * hydro_share, 
             "Q3 hydropower", 
             "Multiply production by hydro share")

check_answer(non_renewable, 1 - renewable, 
             "Q4 non-renewable share", 
             "Subtract renewable share from 1")

check_answer(wind, production * wind_share, 
             "Q5 wind power", 
             "Multiply production by wind share")

```

Output:
```
🔍 Checking all answers...

Q1 surplus is correct
Q2 renewable share is correct
Q3 hydropower is correct
Q4 non-renewable share is correct
Q5 wind power is correct
```

```python
print("\n Final score:")
print(f"{score}/{total} correct")

if score == total:
    print("Perfect!")
elif score >= total * 0.7:
    print("Good job!")
else:
    print("Keep practicing and revise your answers")

```

Output:
```
Final score:
10/10 correct
Perfect!
```

```python

```